Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Rodrigo Erices"
COLLABORATORS = ""
```

# Presenting Uncertainty
## School of Information, University of Michigan

## Week 2: Assignment Overview
Version 1.1
### The objectives for this week are for you to:

- review the concept of standard error
- construct a confidence distribution
- use a confidence distribution to construct a density plot, interval plot, CCDF barplot, and quantile dotplot
- apply one or more of the above techniques to the World Happiness dataset

```python
import numpy as np
import pandas as pd
import seaborn as sns
import matplotlib.pyplot as plt
import altair as alt
import random
from scipy import stats
from scipy.stats import norm
```

# Part 1. Standard error and confidence distributions (14 points)

## 1.1 Bootstrap sampling distribution of phone pickups example

Recall in assignment 1 we used a sample from a hypothetical dataset of counts of how many times students at a school picked up their phones. Let's generate such a hypothetical dataset again. We'll make a sample of 20 students and the number of times each student picked up their phone:

```python
np.random.seed(1234)   # for reproducibility

pickups_df = pd.DataFrame({"pickups": np.random.poisson(lam=40, size=20)})

pickups_df.describe()
```

Output:
```
pickups
count  20.000000
mean   39.350000
std     6.072154
min    27.000000
25%    36.750000
50%    38.500000
75%    43.250000
max    52.000000
```

### Question 1.1.1 (2 points)

Using what you learned in assignment 1, generate a bootstrap sampling distribution of the mean number of pickups and assign it to the variable `bootstrap_means`:

<!-- grade_id: cell-50685f9df244734b, points: 0, grade: false, solution: true -->
```python
B = 5000
np.random.seed(1234) # for reproducibility

bootstrap_means = [np.mean(np.random.choice(pickups_df['pickups'], 20)) for i in range(B)]
```

<!-- grade_id: cell-0881a1b9bd579644, points: 2, grade: true, solution: false -->
```python
# the bootstrap sampling mean should be close to the sample mean
sample_mean = pickups_df['pickups'].mean()
sample_mean_se = np.std(pickups_df['pickups'], ddof=1) / np.sqrt(len(pickups_df))
assert len(bootstrap_means) == 5000, "Bootstrap sampling distribution: length of bootstrap_means should be 5000"
assert np.abs(np.std(bootstrap_means) - sample_mean_se) / sample_mean_se < 0.05, "Bootstrap sampling distribution: SD of boostrap_means does not match SE of sample_pickups"
assert np.abs(np.mean(bootstrap_means) - sample_mean)/sample_mean < 5*1e-3, "Bootstrap sampling distribution: mean of bootstrap_means does not match sample_mean"
```

### Question 1.1.2 (2 points)

Visualize the bootstrap sampling distribution of the mean as a density plot.

Hint: see assignment 1.

<!-- grade_id: cell-1da090945bcdeb0e, points: 0, grade: true, solution: true -->
```python
bootstrap_df = pd.DataFrame({"bootstrap_means": bootstrap_means})

base = alt.Chart(bootstrap_df).transform_density(
    'bootstrap_means',
    as_ = ['bootstrap_means', 'density'])

density_chart = base.mark_area(color = 'green', opacity = 0.3).encode(
    x = 'bootstrap_means:Q',
    y = 'density:Q'
)

mean_chart = base.mark_rule(color = 'red', strokeDash = [1,1]).encode(
    x = alt.X('mean(bootstrap_means)', title = 'mean pickups')
)

density_chart + mean_chart
```

Output:
```
alt.LayerChart(...)
```

## 1.2 Standard error and Normal sampling distribution for mean pickups

The sampling distribution of a mean approaches the shape of a Normal distribution as the sample size increases. Instead of using bootstrapping to derive the distribution, we could instead use the *standard error*, which is an estimate of the standard deviation of the sampling distribution, assuming the sampling distribution is Normal.

The standard error of the mean ($s_\bar{x}$) is calculated using the sample standard deviation ($s$) and the sample size ($n$):

$s_\bar{x} = \frac{s}{\sqrt{n}}$

This calculation can be done as follows:

```python
# manual calculation of standard error
pickups_se = pickups_df['pickups'].std(ddof=1) / np.sqrt(len(pickups_df))
print("Manual calculation of SE:           ", pickups_se)

# or using scipy.stats.sem() (should be the same)
pickups_se = stats.sem(pickups_df['pickups'])
print("scipy.stats.sem() calculation of SE:", pickups_se)

```

Output:
```
Manual calculation of SE:            1.357774882511437
scipy.stats.sem() calculation of SE: 1.357774882511437
```

Given the sample mean pickups ($\bar{x}$ = `pickups_mean`) and the sample mean standard error ($s_\bar{x}$ = `pickups_se`), we can construct a sampling distribution using the standard error. As discussed in lecture, there are three functions we'll want to define in order to construct different uncertainty visualizations: the probability density function (PDF), the cumulative distribution function (CDF), and the quantile function (also called the percentage points function, PPF). For a mean and standard error, these are:

$\begin{align}
\mathrm{density~function~(PDF):~} & f_\mathrm{Normal}\left(x \middle| \bar{x}, s_\bar{x}\right) & \mathtt{stats.norm.pdf()}\\
\mathrm{cumulative~distribution~function~(CDF):~} & F_\mathrm{Normal}\left(x \middle| \bar{x}, s_\bar{x}\right) & \mathtt{stats.norm.cdf()}\\
\mathrm{quantile~function~(PPF):~} & F^{-1}_\mathrm{Normal}\left(p \middle| \bar{x}, s_\bar{x}\right) & \mathtt{stats.norm.ppf()}
\end{align}$

```python
pickups_mean = pickups_df['pickups'].mean()

# generate evenly-spaced numbers covering the range of values of pickups
# that we will calculate the density of the sampling distribution at
x = np.linspace(
    start = pickups_df['pickups'].min(), 
    stop = pickups_df['pickups'].max(),
    num = 1001
)

# this is the density of the sampling distribution: f_Normal(x | x_bar, se)
density = stats.norm.pdf(x, pickups_mean, pickups_se)

df = pd.DataFrame({"Normal sampling distribution of mean pickups": x, "density": density})

pickups_se_chart = alt.Chart(df).mark_line().encode(
    x="Normal sampling distribution of mean pickups",
    y="density"
)

pickups_se_chart
```

Output:
```
alt.Chart(...)
```

## 1.3 t-based confidence distribution for mean pickups

A slight improvement to the sampling distribution above is to use a Student's t confidence distribution. At small sample sizes, the Normal approximation is less accurate; a scaled-and-shifted Student's t distribution will have slightly fatter tails than the Normal distribution. 

To determine how fat the tails are, the t distribution uses a *degrees of freedom* parameter, which for the estimate of a single independent mean of sample size $n$ is equal to $n - 1$. Thus, an improved confidence distribution over the Normal approximation above is:

$\begin{align}
\mathrm{density~function~(PDF):~} & f_\mathrm{t}\left(x \middle| n - 1, \bar{x}, s_\bar{x}\right) & \mathtt{stats.t.pdf()}\\
\mathrm{cumulative~distribution~function~(CDF):~} & F_\mathrm{t}\left(x \middle| n - 1, \bar{x}, s_\bar{x}\right) & \mathtt{stats.t.cdf()}\\
\mathrm{quantile~function~(PPF):~} & F^{-1}_\mathrm{t}\left(p \middle| n - 1, \bar{x}, s_\bar{x}\right) & \mathtt{stats.t.ppf()}
\end{align}$

This confidence distribution is related to the common Student's t test: confidence intervals from this distribution are the same as the confidence intervals you would generate from a t test.

### Question 1.3.1 Visualize the Student's t confidence distribution for the mean (5 points)

Visualize the density of the t-based confidence distribution described above in a similar manner to how the Normal sampling distribution was visualized, but use a different color for the line (this will help you with the next question):


<!-- grade_id: cell-4ded01ddef182e32, points: 0, grade: true, solution: true -->
```python
density_t = stats.t.pdf(x, len(pickups_df) - 1, pickups_mean, pickups_se)

df = pd.DataFrame({"mean pickups": x, "density": density_t})

pickups_se_chart_t = alt.Chart(df).mark_line(color = 'red').encode(
    x="mean pickups",
    y="density"
)

pickups_se_chart_t
```

Output:
```
alt.Chart(...)
```

### Question 1.3.2 Visualize and compare uncertainty distributions (5 points)

Visualize all three distributions you have constructed so far in a single plot. 

Hint: if you have saved your previous plots in three separate variables and made them with different enough encodings that they can be distinguished by eye, this answer can be as simple as `pickups_bootstrap_chart + pickups_se_chart + pickups_t_chart`

<!-- grade_id: cell-95c513c27906f6e5, points: 0, grade: true, solution: true -->
```python
pickups_se_chart_t + pickups_se_chart + density_chart
```

Output:
```
alt.LayerChart(...)
```

Compare these distributions: what do you notice? Write your answer below.

<!-- grade_id: cell-5726ed0c328b15d6, points: 0, grade: true, solution: true -->
The distributions are slightly different. With the red one, having slightly fatter tails as expected. It can be seen that the blue distribution -corresponding to the normal distribution- fits better the density obtained through bootstrapping. This is in accordance to the notion that the normal distribution is a good fit for sufficiently large sample sets, as the bootstrapped sample is quite large. 

# Part 2. Intervals, CDFs, and quantile dotplots (17 points)

Now that we are able to define the density, CDF, and quantile functions to describe our uncertainty in a mean, we can use these to construct various uncertainty visualizations as described in the lectures. For example, we can use the density function to create density plots or gradients; the CDF to create complementary CDF bar plots; and the quantile function to create intervals and quantile dotplots.

## 2.1 Intervals

### Question 2.1.1 Calculate 95% interval (2 points)

Using the quantile function (aka percentage points function) of the Student's t confidence distribution you derived in Part 1, calculate the lower and upper bounds of a 95% confidence interval for mean pickups and assign these bounds to the `pickups_lower` and `pickups_upper` variables. 

(Hint: use `stats.t.ppf()` as the quantile function and see the lecture on intervals.)

<!-- grade_id: cell-9c625ac64cd29f0c, points: 0, grade: false, solution: true -->
```python
pickups_lower = stats.t.ppf(0.025, len(pickups_df) - 1, pickups_mean, pickups_se)
pickups_upper = stats.t.ppf(0.975, len(pickups_df) - 1, pickups_mean, pickups_se)
```

<!-- grade_id: cell-17e176aaae83c5a8, points: 2, grade: true, solution: false -->
```python
dof = len(pickups_df) - 1
m = pickups_df['pickups'].mean()
se = stats.sem(pickups_df['pickups'])
#hiddent tests below
```

### Question 2.1.2 Visualize 95% + 66.666% interval (5 points)

Visualize a combined 95% + 66.666% (i.e. 2/3rds) interval for the mean of pickups, overlaid on top of the data (as ticks). 

Hints: 
- create two different layers, one for the 95% interval and one for the 66.666% interval, then combine them using `alt.layer()` or `+`.
- use `mark_rule(size=XXX).encode(x=YYY, x2=ZZZ)` to create each interval, where you fill in the values of `size` (to set the thickness of the interval), `x` (to set the lower bound), and `x2` (to set the upper bound).

Your output should look like this:

![Plot of 66% and 95% confidence intervals overliad on data](asset/assignment2_intervals.png)

```python
pickups_lower_2nd = stats.t.ppf(1/6, len(pickups_df) - 1, pickups_mean, pickups_se)
pickups_upper_2nd = stats.t.ppf(5/6, len(pickups_df) - 1, pickups_mean, pickups_se)

bounds = pd.DataFrame({'upper': [pickups_upper, pickups_upper_2nd], 'lower': [pickups_lower, pickups_lower_2nd], 'significance': [5, 33]})
```

<!-- grade_id: cell-59b1e61dfc46522c, points: 0, grade: true, solution: true -->
```python
pickups = alt.Chart(pickups_df).mark_rule(color = 'red').encode(x = 'pickups')

intervals = alt.Chart(bounds).mark_rule().encode(
    x = 'lower:Q', 
    x2 = 'upper:Q', 
    size = alt.Size('significance:Q', legend = None)
)

(pickups + intervals)
```

Output:
```
alt.LayerChart(...)
```

## 2.2 Cumulative distribution function

### Question 2.2.1 Visualize a Complementary CDF barplot (5 points)

Using the CDF function (or the CCDF function), plot a CCDF barplot for the Student's t-based confidence distribution. *The barplot should start at 0*.

Hint: start from the code for generating a density plot, above, and (1) adjust the code for generating `x` to ensure it includes 0 and (2) change the function from a PDF to 1 - the CDF.

Your output should look something like this:

![CCDF barplot of mean pickups](asset/assignment2_ccdf.png)

<!-- grade_id: cell-b9ec3dd05162e58b, points: 0, grade: true, solution: true -->
```python
cdf_t = 1 - stats.t.cdf(np.insert(x, 0, 0), len(pickups_df) - 1, pickups_mean, pickups_se)

df = pd.DataFrame({"mean pickups": np.insert(x, 0, 0), "ccdf": cdf_t})

pickups_cdf_chart = alt.Chart(df).mark_area().encode(
    x="mean pickups",
    y="ccdf"
)

pickups_cdf_chart.properties(height = 50)
```

Output:
```
alt.Chart(...)
```

## 2.3 Quantile dotplots

Constructing quantile dotplots involves first generating a small-to-medium number of quantiles from the quantile function (aka percentage points function) of the uncertainty distribution. We do this by projecting evenly-spaced values in probability space back into the data units through the quantile function, then stacking up those values into a dotplot:

<img src="asset/cdf.png" alt="wo graphs stacked on top of each other. The top graph has a y-axis of p less than x, from 0.00 to 1.00 in intervals of 0.25, and the x-axis is x, from 0 to 30. The cumulative distribution function line starts out around (0, 0.00) and stays that way until about (8, 0.00) where it rises up the y axis until it reaches about (18, 1.00) and continues on to 30. The bottom graph has a y-axis of count, from 0.00 to 1.00 in intervals of 0.25, and the x-axis is x, from 0 to 30. Circles are stacked between 8 and 18 on the x-axis, with a peak around (11, 0.50) and is slightly skewed to the right." style="width: 500px;"/>

The first step is to generate the evenly-spaced probabilities. We will create a `ppoints(n)` function that generates `n` evenly-spaced probabilities (not including 0 and 1):

```python
def ppoints(n):
    return (np.arange(n) + 0.5)/n

ppoints(20)
```

Output:
```
array([0.025, 0.075, 0.125, 0.175, 0.225, 0.275, 0.325, 0.375, 0.425,
       0.475, 0.525, 0.575, 0.625, 0.675, 0.725, 0.775, 0.825, 0.875,
       0.925, 0.975])
```

We can then use `ppoints` with the quantile function of the t confidence distribution (`stats.t.ppf`) to construct a quantile dotplot.

Dotplots in Altair are complex to construct, requiring three transformations:

1. `transform_bin` to bin the values
2. `transform_window` to rank the values within each bin (this stacks up all the dots in one bin)
3. `transform_joinaggregate` to calculate the midpoint of each bin using the median of all values in that bin

Putting together `ppoints`, the quantile function of the confidence distribution, and the above steps, we can construct a quantile dotplot.

```python
df = pd.DataFrame({
    "x": stats.t.ppf(ppoints(50), len(pickups_df) - 1, pickups_mean, pickups_se)
})

alt.Chart(df).transform_bin(
    field="x",
    as_="bin",
    bin=alt.BinParams(step=0.75)
).transform_window(
    rank_in_bin="rank()",
    groupby=["bin"]
).transform_joinaggregate(
    bin_midpoint="median(x)",
    groupby=["bin"]
).mark_circle(size=90).encode(
    x="bin_midpoint:Q",
    y="rank_in_bin:Q"
).properties(width=600, height=100)
```

Output:
```
alt.Chart(...)
```

### Question 2.3.1 Construct a 20-dot quantile dotplot for the mean of pickups (5 points)

Construct a 20-dot quantile dotplot for the mean of pickups. Adjust the binning step size, circle size, and/or chart size to achieve a closely-packed quantile dotplot like the 50-dot dotplot above. It is typically also necessary to adjust the step size in the binning (`alt.BinParams(step=XXX)`), the size of the circles (`mark_circle(size=XXX)`) and/or the width and height of the plot (`properties(width=XXX, height=YYY)`) to achieve good-looking dotplots.

<!-- grade_id: cell-ba50d92f593c124d, points: 0, grade: true, solution: true -->
```python
df = pd.DataFrame({
    "x": stats.t.ppf(ppoints(20), len(pickups_df) - 1, pickups_mean, pickups_se)
})

alt.Chart(df).transform_bin(
    field="x",
    as_="bin",
    bin=alt.BinParams(step=1.75)
).transform_window(
    rank_in_bin="rank()",
    groupby=["bin"]
).transform_joinaggregate(
    bin_midpoint="median(x)",
    groupby=["bin"]
).mark_circle(size=100).encode(
    x="bin_midpoint:Q",
    y="rank_in_bin:Q"
).properties(width=350, height=100)
```

Output:
```
alt.Chart(...)
```

## Part 3. Visualize world happiness (10 points)

We will do some visualization of the World Happiness Report dataset from 2015. It is a survey of the state of global happiness, which ranks 155 countries by their happiness levels. Let's explore the happiness score.

First, we'll read in the data and gather the happiness score for all countries:

```python
#Read in 2015 World Happiness Report statistics as a dataframe
happ = pd.read_csv("asset/2015.csv")

#show the top 15 countries:
happ.head(15)
```

Output:
```
Country                           Region  Happiness Rank  \
0     Switzerland                   Western Europe               1   
1         Iceland                   Western Europe               2   
2         Denmark                   Western Europe               3   
3          Norway                   Western Europe               4   
4          Canada                    North America               5   
5         Finland                   Western Europe               6   
6     Netherlands                   Western Europe               7   
7          Sweden                   Western Europe               8   
8     New Zealand        Australia and New Zealand               9   
9       Australia        Australia and New Zealand              10   
10         Israel  Middle East and Northern Africa              11   
11     Costa Rica      Latin America and Caribbean              12   
12        Austria                   Western Europe              13   
13         Mexico      Latin A
```

```python
#some descriptive statistics of the World Happiness Report:
happ.describe()
```

Output:
```
Happiness Rank  Happiness Score  Standard Error  \
count      158.000000       158.000000      158.000000   
mean        79.493671         5.375734        0.047885   
std         45.754363         1.145010        0.017146   
min          1.000000         2.839000        0.018480   
25%         40.250000         4.526000        0.037268   
50%         79.500000         5.232500        0.043940   
75%        118.750000         6.243750        0.052300   
max        158.000000         7.587000        0.136930   

       Economy (GDP per Capita)      Family  Health (Life Expectancy)  \
count                158.000000  158.000000                158.000000   
mean                   0.846137    0.991046                  0.630259   
std                    0.403121    0.272369                  0.247078   
min                    0.000000    0.000000                  0.000000   
25%                    0.545808    0.856823                  0.439185   
50%                    0.910245    1.029510   
```

## 3.1 Happiness Scores

We'll focus on two columns in particular: the `Happiness Score` (an aggregate happiness score for each country) and the `Standard Error`, which is the standard error of the `Happiness Score`: because the data are based on surveys, there is uncertainty in the measurement of the happiness score, and this is quantified using its standard error.

### Question 3.1.1 Visualize the happiness scores and uncertainty of the top 15 countries (5 points)

Use one of the uncertainty visualizations above to visualize the happiness score of each country. You will need to use the Normal approximation to the sampling distribution along with the standard error of the happiness score to do this.

You **MUST NOT** use just a combination of a point and a single interval to visualize the scores. For example, the following chart **IS NOT AN ACCEPTABLE SOLUTION** (but could be a good place to start):

![point estimates and 95% confidence intervals for the happiness score in the top 10 countries according to the World Happiness Report](asset/assignment2_happiness_intervals.png)

But you could construct a chart with a similar layout, but which uses one of the uncertainty encodings you learned above. For example, you might start by constructing a chart with 95% intervals, then extend it to be a chart with 95%+66.666% intervals, which would be an acceptable solution.

Hints:

1. You may want to sort the y axis by happiness score using something like `alt.Y("Country", sort=["Happiness Score"]`
2. You may want to restrict the x axis domain to make it easier to see the uncertainty using something like `alt.X("Happiness Score", scale={"domain":[6,8]})`
3. If you decide to make density plots, quantile dotplots, or CCDF barplots, you need to use the `facet()` function to display multiple subcharts where each one is a single country. If you do that, you may want to set the `"labelAngle"` property on the facet headers so they are horizontal.

**(BONUS)** 2 bonus points if you are able to construct quantile dotplots with well-packed dots.

```python
happ.head()
```

Output:
```
Country          Region  Happiness Rank  Happiness Score  \
0  Switzerland  Western Europe               1            7.587   
1      Iceland  Western Europe               2            7.561   
2      Denmark  Western Europe               3            7.527   
3       Norway  Western Europe               4            7.522   
4       Canada   North America               5            7.427   

   Standard Error  Economy (GDP per Capita)   Family  \
0         0.03411                   1.39651  1.34951   
1         0.04884                   1.30232  1.40223   
2         0.03328                   1.32548  1.36058   
3         0.03880                   1.45900  1.33095   
4         0.03553                   1.32629  1.32261   

   Health (Life Expectancy)  Freedom  Trust (Government Corruption)  \
0                   0.94143  0.66557                        0.41978   
1                   0.94784  0.62877                        0.14145   
2                   0.87464  0.64938                  
```

```python
top15 = happ[happ['Happiness Rank'] <= 15][['Country', 'Happiness Rank', 'Happiness Score', 'Standard Error']]
```

```python
df = pd.DataFrame()

for _, row in top15.iterrows():
    aux = pd.DataFrame()
    aux['x'] = stats.norm.ppf(ppoints(20), row['Happiness Score'], row['Standard Error'])
    aux['Country'] = row['Country']
    aux['Rank'] = row['Happiness Rank']
    df = df.append(aux)
```

<!-- grade_id: cell-3754ddb7718091db, points: 0, grade: true, solution: true -->
```python
alt.Chart(df).transform_bin(
    field="x",
    as_="bin",
    bin=alt.BinParams(step=0.02)
).transform_window(
    rank_in_bin="rank()",
    groupby=["bin", 'Country']
).transform_joinaggregate(
    bin_midpoint="median(x)",
    groupby=["bin", 'Country']
).mark_circle(size=20).encode(
    x = alt.X("bin_midpoint:Q", scale = alt.Scale(domain = [6.8,7.8]), title = 'Happiness Score'),
    y = alt.Y("rank_in_bin:Q", axis = None),
    color = alt.Color('Country', legend = None, scale = alt.Scale(scheme = 'tableau20'))
).properties(
    width=250,
    height=30
).facet(
    row = alt.Row('Country:O',
                  header = alt.Header(labelAngle=0, labelAlign = 'left'), 
                  title = 'Country', 
                  sort = top15.Country.values)
).properties(
    bounds = 'flush',
    title = 'Happiness Score Densities by Country'
).configure_facet(
    spacing = 1
).configure_view(
    stroke = None
).configure_title(
    anchor = 'end')
```

Output:
```
alt.FacetChart(...)
```

### Question 3.1.2 Reflect on the chart you created (5 points)

What are the pros and cons of the uncertainty encoding you used? How might it compare to other options you could have chosen?

<!-- grade_id: cell-1ef207e1b9bef8ef, points: 0, grade: true, solution: true -->
Pros:
- frequency framing helps in avoiding dichotomania, as explicitly showing a distribution throught something countable makes the uncertainty in the score and the ranking harder to avoid
- dotplots allow to do fast counting in tails and body, which can be helpful in better quantify how likely is the score of a country to be closer to a given number 
- also these kind of charts allow to quickly graps where the mode is, which is of course a measure of interest in this case

Cons:
- given size constraints, the shape of the dotplots misleads about the nature of the uncertainty distribution, given our model we know its normal, but in some cases it appears to have 2 modes or be asymetric.
- it does not seem to me that in this context is so useful to be able to do fast counting in either the tails or body, which is one of the main pros of using frequency framing

Comparison:
I think a visualization that allows to cleanly grasp how much overlap can be between the estimated parameter ranges for the countries in the chart would be useful, as that's an interesting comparison to make when understanding ranks by a given metric. A direct density plot, gradient plot, or even the example with point and interval estimates seem more adequate for that purpose. While this kind of encoding also allows such comparison, its not so inmediate and clear.

Please remember to submit both the .pdf and .ipynb formats of your completed notebook.
Please remember to work on your explanations and interpretations!